# Implémenter la classe arbre binaire

> Cours NSI Terminale - Thème 1.

- toc: false 
- badges: true
- comments: false
- categories: [python, NSI, Terminale, Structure_donnees, POO, TP]
- image: images/nsi1.png

Dans ce TP nous allons implémenter une classe permettant de représenter un arbre binaire.

Voici le schéma de la structure envisagée pour la classe **Arbrebin**

![classe Arbrebin](my_icons/classeArbrebin.png)

La propriété **valeur** contiendra la valeur associée au noeud. Les propriétés **gauche** et **droit** seront les sous arbres gauche et droit. Ces deux propriétés seront donc des instances de la classe **Arbrebin**. Si il n'y a pas de sous arbre gauche ou droit, on indiquera la valeur **None** dans les propriétés correspondantes.

La méthode **est_feuille** renverra un bouléen selon que l'objet est une feuille ou non.
La méthode **cree_fils_gauche()** prend en paramètre une valeur et crée un fils à gauche dont la valeur est passée en paramètres.

## Exemple d'utilisation de la classe Arbrebin

En supposant la classe **Arbrebin** créée, voici comment on l'utilise pour créer cet arbre :
![arbre exemple](my_icons/arbre1.png)

```shell
arbre = Arbrebin("A")
sous_arbre_gauche = arbre.cree_fils_gauche("B")
sous_arbre_gauche.cree_fils_gauche("D")
arbre.cree_fils_droit("C")

# Quelques vérifications possibles
print(arbre.est_feuille())
print(arbre.droit.est_feuille())
print(arbre.gauche.valeur)
# Affiche False True B
```

<!-- grade_id: cell-aac985b55a3ec604, points: 0, grade: false, solution: true -->
```python
class Arbrebin():
    
    # la méthode __repr__ définit ce qui sera affiché
    # lorsqu'on tapera l'objet dans Jupyter ou un terminal
    # Ici, on affiche juste al valeur du noeud
    
    def __repr__(self):
        return str(self.valeur)
    
    # Codez ici les méthodes demandées
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# Testez ici les méthodes de votre classe
a = Arbrebin(2)
a
```

<!-- grade_id: cell-eee2cb758a9af1f2, points: 1, grade: true, solution: false -->
```python
# Tester l'exemple de départ

arbre = Arbrebin("A")
sous_arbre_gauche = arbre.cree_fils_gauche("B")
sous_arbre_gauche.cree_fils_gauche("D")
arbre.cree_fils_droit("C")

assert not arbre.est_feuille()
assert arbre.droit.est_feuille()
assert arbre.gauche.valeur == "B"
```

A présent, vous utiliserez la classe Arbrebin et les méthodes que vous avez développées pour représenter l'arbre suivant dans la variable `expr`
![expression](my_icons/expr.png)

Les opérations seront représentées par des chaînes de caractères. Les feuilles seront des entiers.

<!-- grade_id: cell-506da61b6d6152d2, points: 0, grade: false, solution: true -->
```python
expr = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-012eb21015157518, points: 2, grade: true, solution: false -->
```python
# Validation de la réponse
assert expr.valeur == "+"
assert expr.droit.valeur == 1

```

```python

```