# Les ensembles (set)

> Découverte de la structure d'ensembles en Python

- toc: true
- badges: true
- comments: false
- categories: [python, ISN]

Un ensemble en python (***set***) est une structure pouvant contenir plusieurs données, mais contrairement aux listes, ces données sont uniques et non ordonnées. Il n'y a pas de moyen d'accéder à une donnée en particulier en utilisant son numéro d'index.

Les ensembles sont par contre extrèmement efficaces pour la recherche d'un élément : Contrairement aux listes dans lesquelles une recherche impose de parcourir tous les éléments, les ensembles utilisent des techniques d'optimisation (table de hachage) rendant la recherche très performante.

Voici quelques illustrations de l'utilisation des ***set***

```python
# créer un un ensemble

ensemble = {1,5,9,5,1,2,4}
ensemble
```

Comme on peut le voir, les éléments en doubles dans *ensemble* ont été éliminés et l'ordre affiché n'est pas celui dans lequel les éléments ont été saisis.

```python
# essayons quelque chose...

ensemble[3]
```

l'accès aux éléments par indice comme pour les listes n'est pas possible, cela n'a tout simpliement pas de sens.

## Conversion list <-> set

```python
liste = [1,5,9,5,1,2,4]
ensemble = set(liste)
ensemble
```

```python
ensemble = {1, 9, 5, 4, 2}
liste = list(ensemble)
liste
```

## Méthodes sur les ensembles

### ajout et retrait : add et remove

```python
ensemble = {1, 9, 5, 4, 2}
ensemble.add(18)
ensemble.remove(9)
ensemble
```

Attention de bien tester si un élément est dans l'ensemble avant la suppression car sinon...

```python
ensemble.remove(3)
```

et du coup ...
### tester si un élément est présent dans un ensemble : in

```python
3 in ensemble
```

```python
18 in ensemble
```

### Longueur et ensemble vide

```python
# l'ensemble vide est noté {} ou set()
vide = set()
vide.add(3)
vide.remove(3)

# Calculer le nb d'éléments d'un ensemble
len(vide)
```

### Application
Créer une fonction **ensembleCarres** prenant en paramètre un entier $n$ e renvoyant un ensemble contenant les carrés des entiers de 1 à $n$

<!-- grade_id: cell-95cadc4f89194224, points: 0, grade: false, solution: true -->
```python
def ensembleCarres(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-052d63ccbf55abb4, points: 1, grade: true, solution: false -->
```python
ec = ensembleCarres(10)
assert len(ec)==10
assert 64 in ec
```

- Créez une liste *l* de carrés jusqu'à un million. 
- Créez un ensemble *s* de carrés jusqu'à un million. 
- Recherchez si $874466246641$ est un carré

<!-- grade_id: cell-9786c9168deb0bf8, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Comparez les **temps de recherche d'un même nombre** dans l'ensemble et dans la liste. 

<!-- grade_id: cell-ff81f5698da4dc52, points: 1, grade: true, solution: false -->
```python
%%time
assert 874466246641 in s
```

<!-- grade_id: cell-904261d09bae96a6, points: 1, grade: true, solution: false -->
```python
%%time
assert 874466246641 in l
```

### quelques autres méthodes sur les set

    s.isdisjoint(s2)
    s.issubset(s2)
    s.issuperset(s2)
    s <= s2 : inclusion (pareil avec s >= s2).
    s < s2 : inclusion stricte (pareil avec s >= s2).
    set.union(s1, s2, s3) : renvoie la réunion de plusieurs sets.
    set.intersection(s1, s2, s3) : renvoie l'intersection de plusieurs sets

## Exercice

Créer une fonction **ensembleCubes** prenant en paramètre un entier $n$ e renvoyant un ensemble contenant les cubes des entiers de 1 à $n$

<!-- grade_id: cell-c640424ebd05c6d0, points: 0, grade: false, solution: true -->
```python
def ensembleCubes(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e98ff42fb4d66d7b, points: 1, grade: true, solution: false -->
```python
assert 27 in ensembleCubes(10)
```

En déduire en une ligne de python combien de nombres entre 1 et 100 sont à la fois des carrés et des cubes

<!-- grade_id: cell-369bb5f5aa7ee950, points: 0, grade: true, solution: false -->
```python
# Tapez votre ligne dans la cellule ci-dessous
```

```python
# Attention, pas plus d'une ligne de Python !!
```