# Le pendu

> début de jeu de pendu en mode texte.

- toc: true
- badges: true
- comments: false
- categories: [miniprojet, ISN]


## Préparation du jeu

Pour commencer, nous allons créer un dictionnaire en chargeant une liste de mots figurant dans un [fichier texte](liste_mots.txt).

L'import d'un fichier texte dans une liste se fait facilement en Python :

```python
liste_mots = []
with open("liste_mots.txt","r") as fichier:
    liste_mots = fichier.readlines()

len(liste_mots)
```

Output:
```
336529
```

Pour choisir un mot aléatoire dans cette liste, nous utiliserons la fonction choice du module random.

Le mot choisi contient un caractère de fin de ligne parasite (\n) que nous éliminerons grâce au *slicing* : on garde tous les caractères depuis le début jusqu'au dernier, exclu.

Exécutez plusieurs fois la cellule ci-dessous pour voir les différents choix de mots aléatoires.

```python
from random import choice

mot_inconnu = choice(liste_mots)[:-1]

mot_inconnu
```

## Développement des différentes fonctions utiles au jeu

Créez une fonction **masque** prenant en paramètre un mot et renvoyant une chaîne masquée contenant autant de '-' qu'il y a de lettres dans le mot.

Par exemple : *masque('python')* renvoie *'------'*

<!-- grade_id: cell-7c95ec7ae773c1ad, points: 0, grade: false, solution: true -->
```python
def masque(mot):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-f1caa2e890016f80, points: 1, grade: true, solution: false -->
```python
assert masque('python') == '------'
```

Créez une fonction **devine** prenant en paramètres
- le mot inconnu
- le mot masqué
- une lettre
et renvoyant le masque dans lequel la lettre proposée apparaît à toutes les places où celle-ci est présente dans le mot.

Exemple : *devine('python','p----n','t')* renvoie *'p-t--n'*

<!-- grade_id: cell-2e1dbf09c6486dca, points: 0, grade: false, solution: true -->
```python
def devine(mot, mot_masque, lettre):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-be7c5e70c0524f72, points: 1, grade: true, solution: false -->
```python
assert devine('python','p----n','t') == 'p-t--n'
```

## Jouer une partie

Nous disposons à présent des fonctions nécessaires pour jouer une partie. Il ne nous reste donc plus qu'à réaliser la fonction **jouer** qui 
- prend en paramètre une lettre ainsi qu'une liste contenant
    - le mot à deviner
    - le mot masqué (ou partiellement découvert)
    - le nb d'erreurs réalisées
- renvoie la liste actualisée (avec le mot masqué ou le nombre d'erreurs)

**Exemples** : 
- jouer('a', ["python", "------",0]) renvoie ["python", "------",1] car on a une erreur
- jouer('p', ["python", "------",0]) renvoie ["python", "p-----",1] 

<!-- grade_id: cell-e495f659e8e21947, points: 0, grade: false, solution: true -->
```python
def jouer(lettre, etat):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-0dc4e4fb759c2250, points: 0, grade: true, solution: false -->
```python
mot = "python"
mask = masque(mot)
nb_err = 0

assert jouer("a", [mot, mask, nb_err]) == ['python', '------', 1]
assert jouer("p", [mot, mask, nb_err]) == ['python', 'p-----', 0]
```

```python
mot = "python"
etat = [mot, masque(mot),0]
etat = jouer("a", [mot, mask, nb_err])
etat
```

```python
etat = jouer("p", etat)
etat
```

```python
etat = jouer("u", etat)
etat
```

```python
etat = jouer("n", etat)
etat
```

```python
etat = jouer("y", etat)
etat
```

```python
etat = jouer("t", etat)
etat
```

```python
etat = jouer("o", etat)
etat
```

```python
etat = jouer("h", etat)
etat
```

```python

```