# Structures de données linéaires 

> Cours NSI Terminale - Thème 1.

- toc: true 
- badges: true
- comments: false
- categories: [python, NSI, Terminale, Structure_donnees, TP]
- image: images/nsi1.png

## Les piles

On va commencer ce TP par la manipulation des piles, plus faciles à appréhender, et on terminera par la manuipulation des listes chaînées. Rappelons tout d'abord la notion de **piles** répondant à la règle **LIFO** : dernier entré, premier sorti.

![pile](my_icons/pile1.png)

### Description de la structure

Pour stocker les données dans notre pile, nous utiliserons un tableau python (objet list).
Le dernier élément du tableau sera le sommet de la pile. Seul cet élément sera visible.

**Exemple** : Si la pile est représentée en mémoire par le tableau `[2, 3, 5, 8]`, le sommet de la pile sera `8`. Si je dépile le 8, la pile deviendra `[2, 3, 5]` et le sommet de la pile sera 5. Une fois tous les éléments dépilés, la pile sera vide et représentée par `[]`.


### A vous de jouer

Vous allez devoir implémenter les fonctions
- `p_valeur(p)`
    - prend en paramètre une pile `p`
    - renvoie le sommet de la pile ou `None` si la pile est vide
- `p_depile(p)`
    - prend en paramètre une pile `p`
    - dépile le dernier élément saisi
    - renvoie la valeur dépilée ou `None` si la pile est vide
- `p_empile(p, v)`
    - prend en paramètre une pile `p` et une valeur `v`
    - empile la valeur `v`
    - ne renvoie rien

<!-- grade_id: cell-f49fdbbc737964c6, points: 0, grade: false, solution: true -->
```python
def p_valeur(p):
    """- prend en paramètre une pile p
    - renvoie le sommet de la pile
    Exemple : 
    >>> p_valeur([2, 3, 5])
    >>> 5
    >>> p_valeur([])
    >>> None
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-f1f58ee5889bcc38, points: 1, grade: true, solution: false -->
```python
assert p_valeur([]) is None
assert p_valeur([2, 3, 5]) == 5
```

<!-- grade_id: cell-a050d8f552935363, points: 0, grade: false, solution: true -->
```python
def p_depile(p):
    """- prend en paramètre une pile p
    - dépile le dernier élément saisi
    - renvoie le sommet de la pile
    Exemple : 
    >>> p_valeur([2, 3, 5])
    >>> 5
    >>> p_valeur([])
    >>> None
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-9d614672798fb6bb, points: 1, grade: true, solution: false -->
```python
p=[2, 3, 5]
assert p_depile(p) == 5
assert p == [2, 3]
assert p_depile([]) is None
```

<!-- grade_id: cell-105c0d4b0d91a50f, points: 0, grade: false, solution: true -->
```python
def p_empile(p, v):
    """- prend en paramètre une pile p et une valeur v
    - empile la valeur v
    Exemple : 
    >>> p = [2, 3]
    >>> p_empile(p, 5)
    >>> p
    >>> [2, 3, 5]
    """
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-b2e4998e26a60320, points: 1, grade: true, solution: false -->
```python
p = [2, 3]
p_empile(p, 5)
assert p == [2, 3, 5]
```

## Les files

Rappelons la notion de **files** répondant à la règle **FIFO** : premier entré, premier sorti.

![file](my_icons/file1.png)

### Description de la structure

Pour stocker les données dans notre file, nous utiliserons un tableau python (objet list).
- le dernier élément du tableau sera l'avant de la file. Seul cet élément sera visible
- le premier élément du tableau sera l'arrière de la file

**Exemple** : Si la sile est représentée en mémoire par le tableau `[2, 3, 5, 8]`, l'avant de la file sera `8` et l'arrière de la file sera 2. Si on ajoute 1 à l'arrière de la file, celle-ci contiendra `[1, 2, 3, 5, 8]`. Une fois tous les éléments dépilés, la file sera vide et représentée par `[]`.

### A vous de jouer

Vous allez devoir implémenter les fonctions
- `f_valeur(f)`
    - prend en paramètre une file `f`
    - renvoie la valeur à l'avant de la file ou `None` si la file est vide
- `f_defile(f)`
    - prend en paramètre une file `f`
    - défile l'élément situé à l'avant de la file
    - renvoie la valeur défilée ou `None` si la file est vide
- `f_emfile(f, v)`
    - prend en paramètre une file `f` et une valeur `v`
    - emfile la valeur `v` à l'arrière de la file
    - ne renvoie rien

<!-- grade_id: cell-f6e4709b5e9adf4d, points: 0, grade: false, solution: true -->
```python
def f_valeur(f):
    """- prend en paramètre une file f
    - renvoie la valeur à l'avant de la file ou None si la file est vide
    Exemple : 
    >>> f_valeur([2, 3, 5])
    >>> 5
    >>> f_valeur([])
    >>> None
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-cf2a4fa81751c03d, points: 1, grade: true, solution: false -->
```python
assert f_valeur([]) is None
assert f_valeur([2, 3, 5]) == 5
```

<!-- grade_id: cell-06b9232b2cd93301, points: 0, grade: false, solution: true -->
```python
def f_defile(f):
    """- prend en paramètre une file f
    - défile l'élément situé à l'avant de la file
    - renvoie la valeur défilée ou None si la file est vide
    Exemple : 
    >>> f = [2, 3, 5, 8]
    >>> f_defile(f)
    >>> 8
    >>> f
    >>> [2, 3, 5]
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-b6c2cf3fc29a2b12, points: 1, grade: true, solution: false -->
```python
f = [2, 3, 5, 8]
assert f_defile(f) == 8
assert f == [2, 3, 5]
```

<!-- grade_id: cell-2a71af167d61fa9b, points: 0, grade: false, solution: true -->
```python
def f_enfile(f, v):
    """- prend en paramètre une file f et une valeur v
    - enfile la valeur v à l'arrière de la file
    Exemple :
    >>> f = [2, 3, 5, 8]
    >>> f_enfile(f, 1)
    >>> f
    >>> [1, 2, 3, 5, 8]
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-1702ab9f5c717d14, points: 1, grade: true, solution: false -->
```python
f = [2, 3, 5, 8]
f_enfile(f, 1)
assert f == [1, 2, 3, 5, 8]
```

## Les listes chaînées

### Description de la structure

On rappelle que la structure d'une liste chaînée ressemble à ceci :
![Liste chainée](my_icons/listechainee.svg)

L'illustration montre une liste chaînée composée de 4 maillons. Chaque maillon est composé de 2 champs : une valeur et un pointeur vers le maillon suivant.

Pour implémenter une liste chaînée en python nous utiliserons 
- un tableau à 2 éléments décrivant un maillon
    - le premier élément du tableau est la valeur du maillon
    - le second élément est l'indice du maillon suivant (ou None pour le dernier élément)
- un tableau constitué de la liste de tous les maillons

La liste chaînée donnée en illustration pourra donc être représentée ainsi : 

```python
maillons = [[3, 2], [8,None], [5, 1], [2,0]]
premier = 3
```

Vous constaterez que j'ai *volontairement* inscrit dans le tableau `maillons` les éléments de la liste dans le désordre car il n'y a aucune raison que l'ordre des éléments corresponde à l'ordre dans le tableau `maillons` si par exemple on insère en cours de route des éléments au milieu de la liste. 

la variable `premier` contient l'indice du premier élément de la liste, c'est donc le maillon `[2,0]` correspondant à la valeur 2. L'élément suivant sera le maillon à l'indice 0 dans le tableau donc `[3,2]` dont la valeur est 3. Ensuite on accède au maillon d'indice 2 donc `[5,1]` de valeur 5 pour terminer par le maillon `[8,None]` de valeur 8 marquant la fin de la liste. Au final la liste est donc bien $$2 - 3 - 5 - 8$$
On va supposer dans la suite que la variable `maillons` est **globale** et contiendra tous les maillons de toutes les listes avec lesquelles nous travaillerons. Une liste sera donc définie uniquement par l'**indice de son premier élément** dans le tableau ` maillons`.


### A vous de jouer

Vous allez devoir implémenter les fonctions
- `l_premier(p)` 
    - prend en paramètre l'indice du premier élément d'une liste `p`
    - renvoye la valeur du premier élément de la liste 
- `l_valeurs(p)`
    - prend en paramètre l'indice du premier élément d'une liste `p`
    - renvoye un tableau avec toutes les valeurs de la liste
- `l_insere_debut(p, v)`
    - prend en paramètre l'indice du premier élément d'une liste `p` et la valeur à insérer `v`
    - renvoye l'indice du premier élément de la nouvelle liste où la valeur `v` a été insérée au début
- `l_insere_fin(p, v)`
    - prend en paramètre l'indice du premier élément d'une liste `p` et la valeur à insérer `v`
    - insère à la fin de la liste la valeur `v`
    - renvoye `p` car la liste commence par le même élément
- `l_nouveau(v)`
    - prend en paramètre une valeur `v`
    - renvoie l'indice du premier élément de cette liste constituée du seul élément `v`
- `l_scinde(p, i)`
    - prend en paramètre une liste `p` et in indice `i`
    - scinde la liste `p` en deux listes dont la première possède `i` éléments
    - renvoie l'indice du premier élément de la seconde sous liste (la première commence toujours à `p`)
- `l_concat(p1, p2)`
    - prend en paramètre deux listes `p1` et `p2`
    - concatène (met bout à bout) les deux listes
    - renvoie `p1` : début de la liste concaténée.

<!-- grade_id: cell-a8d1b08c34f8179f, points: 0, grade: false, solution: true -->
```python
def l_premier(p):
    """- prend en paramètre l'indice du premier élément d'une liste p
    - renvoye la valeur du premier élément de la liste
    Exemple : maillons = [[3, 2], [8,None], [5, 1], [2,0]]
    >>> l_premier(3)
    >>> 2
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-26ffa2c4a4bbeb20, points: 1, grade: true, solution: false -->
```python
# Testez votre fonction
maillons = [[3, 2], [8,None], [5, 1], [2,0]]
assert l_premier(3) == 2
```

<!-- grade_id: cell-0052c26709ee97b8, points: 0, grade: false, solution: true -->
```python
def l_valeurs(p):
    """
    - prend en paramètre l'indice du premier élément d'une liste p
    - renvoye un tableau avec toutes les valeurs de la liste
    Exemple : maillons = [[3, 2], [8,None], [5, 1], [2,0]]
    >>> l_valeurs(3)
    >>> [2, 3, 5, 8]
    """
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-5aa1a0d5b36d87b6, points: 2, grade: true, solution: false -->
```python
# Testez votre fonction
maillons = [[3, 2], [8,None], [5, 1], [2,0]]
assert l_valeurs(3) == [2, 3, 5, 8]
```

<!-- grade_id: cell-6e8b3a6de3d676c1, points: 0, grade: false, solution: true -->
```python
def l_insere_debut(p, v):
    """-prend en paramètre l'indice du premier élément d'une liste p et la valeur à insérer v
    -renvoye l'indice du premier élément de la nouvelle liste où la valeur v a été insérée au début
    Exemple : maillons = [[3, 2], [8,None], [5, 1], [2,0]]
    >>> p = l_insere_debut(-5)
    >>> l_valeurs(p)
    >>> [-5, 2, 3, 5, 8]
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-d05136a80f6202b8, points: 1, grade: true, solution: false -->
```python
# Testez votre fonction
maillons = [[3, 2], [8,None], [5, 1], [2,0]]
p = l_insere_debut(3,-5)
assert l_valeurs(p) == [-5, 2, 3, 5, 8]
```

<!-- grade_id: cell-3bd6f23d4f965724, points: 0, grade: false, solution: true -->
```python
def l_insere_fin(p, v):
    """-prend en paramètre l'indice du premier élément d'une liste p et la valeur à insérer v
    - insère à la fin de la liste la valeur v
    - renvoye p car la liste commence par le même élément
    Exemple : maillons = [[3, 2], [8,None], [5, 1], [2,0]]
    >>> p = l_insere_fin(3, 11)
    >>> l_valeurs(p)
    >>> [2, 3, 5, 8, 11]
    """
    i = p
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-3f7c56d180ae5e55, points: 2, grade: true, solution: false -->
```python
# Testez votre fonction
maillons = [[3, 2], [8,None], [5, 1], [2,0]]
p = l_insere_fin(3, 11)
assert l_valeurs(p) == [2, 3, 5, 8, 11]
```

<!-- grade_id: cell-113581ada50840b9, points: 0, grade: false, solution: true -->
```python
def l_nouveau(v):
    """- prend en paramètre une valeur v
    - renvoie l'indice du premier élément de cette liste constituée du seul élément v
    Exemple : maillons = [[3, 2], [8,None], [5, 1], [2,0]]
    >>> p = l_nouveau(15)
    >>> l_valeurs(p)
    >>> [ 15 ]
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-11ad46a872422284, points: 1, grade: true, solution: false -->
```python
# Testez votre fonction
maillons = [[3, 2], [8,None], [5, 1], [2,0]]
p = l_nouveau(15)
assert l_valeurs(p) == [ 15 ] 
```

<!-- grade_id: cell-70c3e5cb2232ae7f, points: 0, grade: false, solution: true -->
```python
def l_scinde(p, i):
    """- prend en paramètre une liste p et in indice i
    - scinde la liste p en deux listes dont la première possède i éléments
    - renvoie l'indice du premier élément de la seconde sous liste (la première commence toujours à p)
    Exemple : maillons = [[3, 2], [8,None], [5, 1], [2,0]]
    >>> p1 = 3
    >>> p2 = l_scinde(p1, 2)
    >>> l_valeurs(p1)
    >>> [ 2, 3 ]
    >>> l_valeurs(p2)
    >>> [ 5, 8 ]
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-8420a311905dc42f, points: 2, grade: true, solution: false -->
```python
maillons = [[3, 2], [8,None], [5, 1], [2,0]]
p1 = 3
p2 = l_scinde(p1, 2)
assert l_valeurs(p1) == [2, 3]
assert l_valeurs(p2) == [5, 8]
p2
```

<!-- grade_id: cell-64ba676ccd832169, points: 0, grade: false, solution: true -->
```python
def l_concat(p1, p2) :
    """- prend en paramètre deux listes p1 et p2
    - concatène (met bout à bout) les deux listes
    - renvoie p1 : début de la liste concaténée.
    Exemple : [[3, None], [8, None], [5, 1], [2, 0]]
    >>> l_concat(3, 2)
    >>> l_valeurs(3)
    >>> [ 2, 3, 5, 8]
    """
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-62fcce8c5a8ad05c, points: 2, grade: true, solution: false -->
```python
maillons = [[3, None], [8, None], [5, 1], [2, 0]]
assert l_concat(3, 2) == 3
assert l_valeurs(3) == [2, 3, 5, 8]
```

```python

```